# 12. Exercises: Object-Oriented Design

## Exercise 12.1

Create a class to represent vectors of arbitrary length and which is initialised with a list of values, e.g.:

```python
x = MyVector([0, 2, 4])
```

Equip the class with methods that:

1. Return the length of the vector (method name `size`)
2. Compute the norm of the vector $\sqrt{x \cdot x}$ (method name `norm`)
3. Compute the dot product of the vector with another vector (method name `dot`)

Test your implementation using two vectors of length 3. To help you get started, a skeleton of the class is provided below. Don't forget to use `self` where necessary.

```python
import math

class MyVector:
    """A vector object that can return its size and norm, and can compute the dot product with another vector."""

    def __init__(self, x):
        self.x = x

    def size(self):
        """Return length of vector."""
        return len(self.x)

    def __getitem__(self, index):
        """This allows access by index, e.g. y[2]."""
        return self.x[index]

    def norm(self):
        """Return norm of vector."""
        dot_product = self.dot(self)
        return math.sqrt(dot_product)

    def dot(self, other):
        """Return dot product of vector with another vector."""
        assert(self.size() == other.size())
        total = 0
        for x, y in zip(self.x, other):
            total += x * y
        return total
```

<!-- grade_id: cell-b601b5cc8cf8ab44, points: 0, grade: true, solution: false -->
```python
# Create two vectors
u = MyVector([1, 1, 2])
v = MyVector([2, 1, 1])

assert u.size() == 3
assert round(u.norm() - 2.449489742783178) == 0.0
assert round(u.dot(v) - 5.0, 10) == 0.0
```

## Exercise 12.2

1. Create a class for holding a student record entry. It should have the following attributes:
   - Surname
   - Forename
   - Birth year
   - Graduation year
   - University
   - Username (optional field)

2. Equip your class with the method '`age`' that returns the age of the student in whole years

3. Equip your class with the method '`__repr__`' such using `print` on a student record displays with the format

       Surname: Bloggs, Forename: Andrea, University: Cambridge

4. Equip your class with the method `__lt__(self, other)` so that a list of record entries can be sorted by (surname, forename) using the Python built-in sort function.

   Create a list of entries and test the sorting. Make sure you have two entries with the same surname.

> Recall that the methods starting with `__`, e.g. `__lt__` and `__repr__`, should **not** be called directly.
> Python with map them to other operations, e.g. `__repr__` is called when using `print`, and `__lt__` is called when using `<`. These functions must have a return value.

*Hint:* To get the current year:

```python
import datetime
year = datetime.date.today().year
print(year)
```

<!-- grade_id: cell-6e857f6089fe0ba8, points: 0, grade: false, solution: true -->
```python
class StudentEntry:
    def __init__(self, surname, forename, birth_year, year, university, username=None):
        self.surname = surname
        self.forename = forename
        self.birth_year = birth_year
        self.year = year
        self.university = university
        self.username = username

    def __repr__(self):
        return (
            f"Surname: {self.surname}, "
            f"Forename: {self.forename}, "
            f"University: {self.university}"
        )

    def __lt__(self, other):
        if (self.surname < other.surname or
            self.surname == other.surname and self.forename < other.forename):
            return True
        else:
            return False

    def age(self):
        return datetime.date.today().year - self.birth_year
```

<!-- grade_id: cell-ce8aafb3e627dd86, points: 0, grade: true, solution: false -->
```python
s0 = StudentEntry("Bloggs", "Andrea", 1996, 1, "Cambridge", "ab1001")
s1 = StudentEntry("Reali", "John", 1997, 1, "UCL")
s2 = StudentEntry("Bacon", "Kevin", 1996, 1, "QMUL")
s3 = StudentEntry("Bacon", "Alexander", 1996, 1, "Southampton")

assert s0 < s1
assert s0 > s2
assert s3 < s2
assert s0.age() ==  datetime.date.today().year - 1996
assert s1.age() ==  datetime.date.today().year - 1997
assert str(s1) == "Surname: Reali, Forename: John, University: UCL"

# Test sorting
s = [s0, s1, s2, s3]
s.sort()
for earlier, later in zip(s, s[1:]):
    assert earlier.surname <= later.surname
    if earlier.surname == later.surname:
        assert earlier.forename <= later.forename
```