# 5. Exercises: Library functions

## Exercise 05.1

- Using the ```randint``` function from the ```random``` module (https://docs.python.org/3/library/random.html#random.randint), develop a function ```dice_roll``` that emulates the roll of a dice with $n$ sides.

- For $n=6$, devise and implement a test to check that it is a fair dice.

#### (a) Dice roll code:

<!-- grade_id: cell-70a83466894475ee, points: 0, grade: false, solution: true -->
```python
import random

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0970743b0c035f78, points: 0, grade: true, solution: false -->
```python
for i in range(1, 20):
    for j in range(100):
        value = dice_roll(i)
        assert value >= 1 and value <= i
```

#### (b) Test for fairness

<!-- grade_id: cell-0826760c62c21bff, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 05.2

For devices with limited memory, data compression can be important. Data compression is a field of its own, but with libraries we can compress (and uncompress) data easily without being expert in the details.

Below is a program code for compressing a passage from Hamlet, by Shakespeare.

```python
# Import the compression module
import zlib

# Create a string that we wish to compress
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""

# Convert Python string to bytes and check type
text_bytes = text.encode("utf-8")
print(type(text_bytes))

# Get number of bytes used to store string
print("Number of bytes for uncompressed string:", len(text_bytes))

# Compress string and get number of byes used for compressed string
text_comp = zlib.compress(text_bytes)
print("Number of bytes for compressed string:", len(text_comp))

# Display the compression efficiency
print("Compression efficiency: ", len(text_comp)/len(text_bytes))

# Decompress the string
text_decomp = zlib.decompress(text_comp)

# Check that original and decompressed string are the same (more on assert)
if text != text_decomp.decode("utf-8"):
    print("Problem: original and decompressed string differ.")
```

Using the above as a guide, examine the compression efficiency of:

1. Compressing one large string made up of the the passage by Shakespeare repeated 100 times; and
2. Compressing a random string of the same length as the repeated Shakespeare passage.

```python
import random
import string

# Shakespeare string
text = """
Welcome, dear Rosencrantz and Guildenstern!
Moreover that we much did long to see you,
The need we have to use you did provoke
Our hasty sending. Something have you heard
Of Hamlet's transformation; so call it,
Sith nor the exterior nor the inward man
Resembles that it was. What it should be,
More than his father's death, that thus hath put him
So much from the understanding of himself,
I cannot dream of: I entreat you both,
That, being of so young days brought up with him,
And sith so neighbour'd to his youth and havior,
That you vouchsafe your rest here in our court
Some little time: so by your companies
To draw him on to pleasures, and to gather,
So much as from occasion you may glean,
Whether aught, to us unknown, afflicts him thus,
That, open'd, lies within our remedy."""

def random_string(N):
    """Generates a random string of length 'N'"""
    return ''.join([random.choice(string.ascii_letters + string.digits) for n in range(N)])

print(random_string(8))
```

```python
"""Import the necessary modules"""
import random
import string
import zlib

"""Repeat the Shakespeare string 100 times and compress"""
# YOUR CODE HERE
raise NotImplementedError()

"""Create random string and compress"""
# YOUR CODE HERE
raise NotImplementedError()

"""Compare compression efficiency"""
# YOUR CODE HERE
raise NotImplementedError()
```