# 7. Exercises: Numerical computation

```python
import numpy as np
```

## Exercise 07.1

Create two very long NumPy arrays ```x``` and ```y``` and sum the arrays using:

1. The NumPy addition syntax, ```z = x + y```; and
2. A ```for``` loop that computes the sum entry-by-entry

Compare the time required for the two approaches for vectors of different lengths (use a very long vector for the timing). The values of the array entries are not important for this test. Use ```%time``` to report the time.

#### (1) Add two vectors using built-in addition operator:

<!-- grade_id: cell-2c856e54f7c3340e, points: 0, grade: false, solution: true -->
```python
x = np.random.rand(10000000)
y = np.random.rand(10000000)

z1 = [0] * len(x)

%time z1 = np.add(x, y)
```

#### (2) Add two vectors using own implementation:

<!-- grade_id: cell-f34614f9f0068cc4, points: 0, grade: false, solution: true -->
```python
z2 = [0] * len(x)

%time z2 = [a + b for a,b in zip(x, y)]
```

```python
z3 = [0] * len(x)
```

```python
%%time
for i in range(len(x)):
    z3[i] = x[i] + y[i]
```

### Optional extension: just-in-time (JIT) compilation

You will see a large difference in the time required between your NumPy and 'plain' Python implementations. This is due to Python being an *interpreted* language as opposed to a *compiled* language. A way to speed up plain Python implementions is to convert the interpreted Python code into compiled code. A tool for doing this is [Numba](https://numba.pydata.org/).

Below is an example using Numba and JIT to accelerate a computation:

```python
!pip -q install numba
import numba
import math

def compute_sine_native(x):
    """Computes the sine of each number in a list."""
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z

@numba.jit
def compute_sine_jit(x):
    """Computes the sine of each number in a list using JIT."""
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z

x = np.ones(10000000)
%time z = compute_sine_native(x)
compute_sine_jit(x)
%time z = compute_sine_jit(x)
```

**Task:** Test if Numba can be used to accelerate your implementation that uses indexing to sum two arrays, and by how much.

```python
@numba.jit
def add_vectors(x, y):
    """Adds twp vectors of arbitrary length using JIT."""
    z = [0] * len(x)
    z = [a + b for a,b in zip(x, y)]
    return z

x = np.random.rand(10000000)
y = np.random.rand(10000000)

%time z = add_vectors(x, y)
```

## Exercise 07.2

Anonymised scores (out of 60) for an examination are stored in a NumPy array. Write:

1. A function that takes a NumPy array of the raw scores and returns the scores as percentages, sorted from lowest to highest (try using ```scores.sort()```, where ```scores``` is a NumPy array holding the scores).

2. A function that returns the maximum, minimum and mean of the raw scores as a dictionary with the keys ```min```, ```max``` and ```mean```. Use the NumPy array methods (functions) ```min()```, ```max()``` and ```mean()``` to do the computation, e.g. ```max = scores.max()```.

   Design your function for the min, max and mean to optionally exclude the highest and lowest scores from the computation of the min, max and mean.

   *Hint:* sort the array of scores and use array slicing to exclude the first and the last entries.

<!-- grade_id: cell-169ebae60810c6be, points: 0, grade: false, solution: true -->
```python
def to_percentage_and_sort(scores):
    """Turns a list of scores out of 60 into a sorted
       list of percentages."""
    scores.sort()
    percentages = np.divide(scores, 0.6)
    return percentages

def statistics(scores, exclude=False):
    """Provide statistics on a list of scores"""
    if exclude:
        scores.sort()
        sorted_scores = scores[1:-1]
    else:
        scores.sort()
        sorted_scores = scores[:]

    min = sorted_scores.min()
    max = sorted_scores.max()
    mean = sorted_scores.mean()
    stats = {"min": min,
            "max": max,
            "mean": mean,
            }

    return stats
```

<!-- grade_id: cell-af0b6fd8a3cadb1a, points: 0, grade: true, solution: false -->
```python
scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
assert np.isclose(to_percentage_and_sort(scores), [ 13.0, 40.0, 58.33333333,  70.0, 96.66666667]).all()

s0 = statistics(scores)
assert round(s0["min"] - 7.8, 10) == 0.0
assert round(s0["mean"] - 33.36, 10) == 0.0
assert round(s0["max"] - 58.0, 10) == 0.0

s1 = statistics(scores, True)
assert round(s1["min"] - 24.0, 10) == 0.0
assert round(s1["mean"] - 33.666666666666666667, 10) == 0.0
assert round(s1["max"] - 42.0, 10) == 0.0
```

## Exercise 07.3

For the two-dimensional array

```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
print(A)
```

use array slicing for the below operations, printing the results to the screen to check. Try to use array slicing such that your code would still work if the dimensions of ```A``` were enlarged.

#### 1. Extract the third column as a 1D array

<!-- grade_id: cell-dfee6b0ed9343682, points: 0, grade: false, solution: true -->
```python
third_column = np.zeros(len(A))
for i, row in enumerate(A):
    third_column[i] = row[2]

print(third_column)
```

#### 2. Extract the first two rows as a 2D sub-array

<!-- grade_id: cell-7bf2f9a8c67029f8, points: 0, grade: true, solution: true -->
```python
top_two_rows = A[:2]

print(top_two_rows)
```

#### 3.  Extract the bottom-right $2 \times 2$ block as a 2D sub-array

<!-- grade_id: cell-5206fea47d246222, points: 0, grade: false, solution: true -->
```python
bottom_2x2_block = np.array([np.zeros(2),
                             np.zeros(2)])
for i in range(0,2):
    for j in range(0,2):
        bottom_2x2_block[i][j] = A[i+2][j+2]

print(bottom_2x2_block)
```

#### 4. Sum the last column

<!-- grade_id: cell-34f74988e47b9f87, points: 0, grade: false, solution: true -->
```python
sum = 0
for i, row in enumerate(A):
    sum += row[3]

print(sum)
```

#### Compute transpose

Compute the transpose of ```A``` (search online to find the function/syntax to do this).

<!-- grade_id: cell-dd2999d2da8070f6, points: 0, grade: false, solution: true -->
```python
AT = A.transpose()

print(AT)
```

## Exercise 07.4

In a previous exercise you implemented the bisection algorithm to find approximate roots of a mathematical function. Use the SciPy bisection function ```optimize.bisect``` (https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.bisect.html) to find roots of the mathematical function that was used in the previous exercise. Compare the results computed by SciPy and your program from the earlier exercise, and compare the computational time (using ```%time```).

<!-- grade_id: cell-b6580accbcb3c2da, points: 0, grade: false, solution: true -->
```python
from scipy import optimize

def fn(x):
    y = x**3 - 6*x**2 + 4*x + 12
    return y

# Set parameters
tol = 5.0e-13
error = tol + 1.0
max_iter = 100000

# Initial end points
x0 = 3.0
x1 = 6.0
```

```python
%%time

x = optimize.bisect(fn, x0, x1, xtol=tol, maxiter=max_iter)

print(x)
```

```python
%%time

# Iterate until tolerance is met
counter = 0
while error > tol:
    counter += 1 # increment counter

    # Compute midpoint
    x_mid = (x0 + x1)/2

    # Evaluate function at left end-point and at midpoint
    f0 = fn(x0)
    f = fn(x_mid)

    error = abs(f)

    if f0 * f < 0:
        x0, x1 = x0, x_mid
    else:
        x0, x1 = x_mid, x1

    # Guard against an infinite loop
    if counter > max_iter:
        print("Oops, iteration count is very large. Breaking out of while loop.")
        break

    # print(counter, x_mid, error)

print(x_mid)
```