# 6. Exercises: Data structures

## Exercise 06.1

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high.

Use appropriate data structures, e.g. ```list```s, ```tuple```s, ```dict```s, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area_lists(v0, v1, v2):
    """Calculate area of a triangle, given the vertices, using lists for each coordinate."""
    """Each vertex must be a list of form [x_coordinate, y_coordinate]"""
    ans = abs(v0[0]*(v1[1]-v2[1]) + v1[0]*(v2[1]-v0[1])+v2[0]*(v0[1]-v1[1]))/2
    return ans

def area_tuples(v0, v1, v2):
    """Calculate area of a triangle, given the vertices, using tuples for each coordinate."""
    """Each vertex must be a tuple of form (x_coordinate, y_coordinate)"""
    ans = abs(v0[0]*(v1[1]-v2[1]) + v1[0]*(v2[1]-v0[1])+v2[0]*(v0[1]-v1[1]))/2
    return ans

def area_dicts(v0, v1, v2):
    """Calculate area of a triangle, given the vertices, using dicts for each coordinate."""
    """Each vertex must be a dict of form {"x": x_coordinate, "y": y_coordinate}"""
    ans = abs(v0["x"]*(v1["y"]-v2["y"]) + v1["x"]*(v2["y"]-v0["y"])+v2["x"]*(v0["y"]-v1["y"]))/2
    return ans
```

### Optional extension

Have a go at implementing your own tests, using the data structures that your function requires.

*Hint*: Look at the tests used in Exercise 04.

```python
v0 = [0.0, 0.0]
v1 = [0.0, 2.0]
v2 = [3.0, 0.0]
A = area_lists(v0, v1, v2)
assert round(A - 3.0, 10) == 0.0
```

## Exercise 06.2

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by

$$ A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
def area(list_of_vertex_tuples):
    """Computes area of a simple n sided polygon, given a vertex coordinates, given as a list of (x,y) tuples."""
    A = 0
    for i in range(0, len(list_of_vertex_tuples)):
        v_i = list_of_vertex_tuples[i]
        v_ip1 = list_of_vertex_tuples[(i+1) % len(list_of_vertex_tuples)]
        A += 0.5 * abs(v_i[0] * v_ip1[1] - v_ip1[0] * v_i[1])

    return A
```

## Exercise 06.3

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

*Hint: You can create a list of zeros of length ```n``` by*

```python
    z = [0]*n
```

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    if len(x) != len(y):
        print("Uh oh... the two vectors have different sizes ({} and {})".format(len(x), len(y)))
    else:
        ans = [0] * len(x)
        for i in range(len(x)):
            ans[i] = x[i] + y[i]

        return ans
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced)

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip). Store the answer in a variable called ```d```.

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
d = [i+j for i, j in zip(a, b)]
```

```python
d = sum_vector(a, b)
assert d == [-2, 11.3, -20, 8]
```

## Exercise 06.4

Create a dictionary that maps names of players on sports teams (the key) to the team they play for, for at least 5 athletes. From the dictionary, produce and print:

1. A dictionary from player to team; and
1. A list of player names sorted into alphabetical order.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
players = {
    "Justin Fields": "Bears",
    "Patrick Mahomes": "Chiefs",
    "Bruce Irvin": "Bears",
    "Tom Brady": "Buccaneers",
    "Joe Burrow": "Bengals"
}

names = sorted([name for name in players.keys()])
print(names)
```

### Optional extension

Create a dictionary that maps the names of players on sports teams (the key) to dictionaries of:

- The team they play for
- Year of birth
- Any numerical statistic e.g. goals scored, races won, etc.

for at least 5 athletes. Using this dictionary,

1. Find the athlete with the greatest *statistic* and print their year of birth.
2. Find the oldest athlete, and print the numerical statistic and their year of birth.

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
players = {
    "Justin Fields": {"Team": "Bears", "Year": 1999, "Touchdowns": 9},
    "Patrick Mahomes": {"Team": "Chiefs", "Year": 1995, "Touchdowns": 145},
    "Russell Wilson": {"Team":"Seahawks", "Year": 1988, "Touchdowns": 307},
    "Tom Brady": {"Team": "Buccaneers", "Year": 1977, "Touchdowns": 645},
    "Joe Burrow": {"Team": "Bengals", "Year": 1996, "Touchdowns": 48},
}

max = 0
for p in players:
    player = players[p]
    if player["Touchdowns"] > max:
        year_of_birth = player["Year"]

print(year_of_birth)

year = 2022
for p in players:
    player = players[p]
    if player["Year"] < year:
        touchdowns = player["Touchdowns"]
        year_of_birth = player["Year"]

print(touchdowns, year_of_birth)
```